# Getting Started

Usually, getting your environment set up is one of the most time consuming and frustrating parts of development. If you are seeing this running locally on your computer, congratulations!! You're past most of the hardship.

During class, we'll be using [colab](https://g.co/colab) for our exercises but I urge you to [install Python](https://www.python.org/getit/) on your computers, [pipenv](https://docs.pipenv.org/) and set up [Jupyter](http://jupyter.org/). If you need help with this, please email me or come to office hours. 

In this exercise we're going to briefly go over Jupyter Notebooks, Markdown, Python, data visualization code and the machine learning workflow. There's a lot to cover so let's get started.

## Jupyter

Right now, you're looking at a Jupyter notebook. Jupyter notebooks allow you to interactively run code while maintaining variables' states in memory. Notebooks are extremely useful for getting your ideas tested and quickly iterating on code. Once you've verified and validated your ideas, you should place any function definitions and workflow procedures in `.py` files. 

Notebooks have multiple cells, each cell has a code type and contents as well as output or display. This cell you're looking at is a Markdown cell. You can double click it to see (and edit) its contents. 

If you haven't used Markdown before, this [cheatsheet](https://github.com/adam-p/markdown-here/wiki/Markdown-Cheatsheet) is very useful. A portion of its contents are copied in the next cell for you to edit and experiment with.

## Markdown Cheat Sheet

**If you already know markdown, you can skip this whole cell.**

<-- Cheatsheet contents -->

### Headings

# H1
## H2
### H3
#### H4
##### H5
###### H6

Alternatively, for H1 and H2, an underline-ish style:

Alt-H1
======

Alt-H2
------

### Emphasis

Emphasis, aka italics, with *asterisks* or _underscores_.

Strong emphasis, aka bold, with **asterisks** or __underscores__.

Combined emphasis with **asterisks and _underscores_**.

Strikethrough uses two tildes. ~~Scratch this.~~

### Lists

1. First ordered list item
2. Another item
⋅⋅* Unordered sub-list. 
1. Actual numbers don't matter, just that it's a number
⋅⋅1. Ordered sub-list
4. And another item.

⋅⋅⋅You can have properly indented paragraphs within list items. Notice the blank line above, and the leading spaces (at least one, but we'll use three here to also align the raw Markdown).

⋅⋅⋅To have a line break without a paragraph, you will need to use two trailing spaces.⋅⋅
⋅⋅⋅Note that this line is separate, but within the same paragraph.⋅⋅
⋅⋅⋅(This is contrary to the typical GFM line break behaviour, where trailing spaces are not required.)

* Unordered list can use asterisks
- Or minuses
+ Or pluses

### Links


[I'm an inline-style link](https://www.google.com)

[I'm an inline-style link with title](https://www.google.com "Google's Homepage")


### Images

Here's a logo (hover to see the title text):

Inline-style: 
![alt text](https://github.com/adam-p/markdown-here/raw/master/src/common/images/icon48.png "Logo Title Text 1")

Reference-style: 
![alt text][logo]


### Code

Inline `code` has `back-ticks around` it.

```javascript
var s = "JavaScript syntax highlighting";
alert(s);
```
 
```python
s = "Python syntax highlighting"
print s
```
 
```
No language indicated, so no syntax highlighting. 
But let's throw in a <b>tag</b>.
```

### Tables

Colons can be used to align columns.

| Tables        | Are           | Cool  |
|---------------|---------------|-------|
|1|2|3|
|4|5|6|
|7|8|9|

There must be at least 3 dashes separating each header cell.
The outer pipes (|) are optional, and you don't need to make the 
raw Markdown line up prettily.


Jupyter organizes content in cells. Each cell has a type such as Code or Markdown and can be run to produce an output or render itself. 

Jupyter also has some cool features that allow us to apply commands to a cell. These are called [magic commands](https://ipython.readthedocs.io/en/stable/interactive/magics.html). 

We use magic commands by typing `%` for line magic and `%%` for cell magic. Let's try the `%%time` magic command.

```python
import time
```

```python
%%time
time.sleep(10) # sleep for 10 seconds
```

## Python

Let's go over some basic Python.

We can do some basic maths as follows:

```python
1 + 1
```

```python
x = 5
y = 2
x + y, x / y, x // y, x **y, x % y
```

We can also do some simple logic flow:

```python
score = 95
if score > 90:
    letterGrade = "A"
elif score > 80:
    letterGrade = "B"
elif score > 70:
    letterGrade = "C"
else:
    letterGrade = "F"

print(f"That score gets a letter grade of {letterGrade}.")
```

This is how we iterate in a loop:

```python
for i in range(1,11):
    if i == 1:
        print("This is the 1st time.")
    elif i == 2:
        print("This is the 2nd time.")
    elif i == 3:
        print("This is the 3rd time")
    else:
        print(f"This is the {i}th time")
```

List comprehensions are a very powerful and useful tool in Python, here are a couple examples:

```python
[i+1 for i in range(10)]
```

```python
[i**2 for i in range(10)]
```

We can also define some functions in Python. Here's an example function that gets the square root of sum of the squares in a list. Does this sound familiar?

```python
def get_metric(vector):
    """Get the square root of the sum of the squares in an iterable.

    Args:
        vector (iterable): the data to calculate the metric of
    
    Returns:
        float
    """
    return sum([i ** 2 for i in vector])**0.5
```

```python
get_metric([3,4])
```

That metric was the $\ell_2$-norm. Can you write a function that takes a vector and value for p and determines the $\ell_p$-norm of a vector? 

<!-- grade_id: get_lp, points: 0, grade: false, solution: true -->
```python
def get_lp(vector, p):
    """Get the $\ell_p$-norm of data

    Args:
        vector (iterable): the data to calculate the metric of
        p (integer): the value of p in the $\
    
    Returns:
        float
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

```python
import pytest
```

<!-- grade_id: test_get_lp, points: 10, grade: true, solution: false -->
```python
assert get_lp([3,4],2) == 5
assert get_lp([3,4,5],2) == pytest.approx(7,0.1)
assert get_lp([10], 5) == pytest.approx(10)
```

We can also write functions in shorthand using `lambda` this creates an anonymous function. This is useful when we're using maps and filters.

```python
lambda x : x**2
```

```python
example_map = map(lambda x: x**3,[1,2,3,4,5])
list(example_map)
```

```python
example_filter = filter(lambda x: x%2 == 1, [1,2,3,4,5])
list(example_filter)
```

Python has a great ecosystem of libraries available and we will be using them for almost all the code we write. Here's how to import code in Python. I recommend reading up on the intricacies of imports. [1](https://stackoverflow.com/a/710603) [2](https://stackoverflow.com/a/19185936)

You're probably going to see the below imports and couple lines of code very frequently at the top of most files we write. First we import the libraries we often use; [numpy](https://www.numpy.org/), [pandas](https://pandas.pydata.org/), [matplotlib](https://matplotlib.org/), and [scikit-learn](https://scikit-learn.org/). The last two lines are for setting plots to show in the notebook directly and not open up in another window and to select the styleguide for the plots we make to ggplot. You can select another style from the [available options](https://matplotlib.org/gallery/style_sheets/style_sheets_reference.html) if you like.


```python
import numpy as np
import pandas as pd
import matplotlib
import matplotlib.pyplot as plt
import sklearn as sk

%matplotlib inline
plt.style.use("ggplot")
```

## Data Visualization

Let's go over some data visualization basics. To start, we'll need some data. We'll create data from a function with some noise and see if we can fit a polynomial to it. Parts of the following are extracted from the following sklearn tutorials. [1](http://scikit-learn.org/stable/auto_examples/model_selection/plot_underfitting_overfitting.html) [2](http://scikit-learn.org/stable/auto_examples/model_selection/plot_validation_curve.html)

```python
def underlying_pattern(X, noise=0.1):
    """Create a function to fit
    
    Args:
        X numpy.array : array of x values
        
    Returns:
        numpy.array : array of output values
    """
    return np.cos(1.5 * np.pi * X) + np.random.randn(X.shape[0]) * noise
```

You should always set a seed value so that you get reproducable results.

```python
np.random.seed(0)
```

```python
n = 20
X = np.sort(np.random.rand(n))
y = underlying_pattern(X)
```

```python
plt.plot(X, y, marker="o", linestyle="None")
# plt.scatter(X, y) # This might be easier to remember, both functions accomplish the same thing. Try switching to the commented one.
plt.xlabel("x")
plt.ylabel("y")
plt.title("Sample Data")
plt.xlim((0, 1))
plt.ylim((-2, 2))
plt.show()
```

## Machine Learning Workflow

Now that we have some data, let's try to fit a machine learning model, the one we already covered, polynomial regression, to the data. For now, you'll get all the code you need to run this experiment.

```python
from sklearn.model_selection import train_test_split
from sklearn.pipeline import Pipeline
from sklearn.preprocessing import PolynomialFeatures
from sklearn.linear_model import LinearRegression
```

```python
X_train, X_test, y_train, y_test = train_test_split(X, y, random_state=0, test_size=0.2)
X_model, X_valid, y_model, y_valid = train_test_split(X_train, y_train, random_state=0, test_size=0.2)

print(f"All Data:        {len(X)} points")
print(f"Training data:   {len(X_train)} points")
print(f"Testing data:    {len(X_test)} points")
print(f"Modeling data:   {len(X_model)} points")
print(f"Validation data: {len(X_valid)} points")
```

```python
plt.plot(X_model, y_model, marker="o", linestyle="None")
# plt.scatter(X, y) # This might be easier to remember, both functions accomplish the same thing. Try switching the commented one.
plt.xlabel("x")
plt.ylabel("y")
plt.title("Model Data")
plt.xlim((0, 1))
plt.ylim((-2, 2))
plt.show()
```

Here you can train a model by manually adjusting the hyperparameter, in this case the polynomial order.

```python
import math
```

```python
## Change this value and check out the resulting plot
order = 1


## Don't worry about the below code for now
xMin = math.floor(X_model.min())
xMax = math.ceil(X_model.min())
polynomial_features = PolynomialFeatures(degree=order, include_bias=False)
linear_regression = LinearRegression()
pipeline = Pipeline([("polynomial_features", polynomial_features),
                         ("linear_regression", linear_regression)])
pipeline.fit(X_model[:, np.newaxis], y_model)

plt.plot(X_model, y_model, marker="o", linestyle="None")
plt.plot(np.linspace(xMin,xMax), pipeline.predict(np.linspace(xMin,xMax)[:, np.newaxis]))
plt.xlabel("x")
plt.ylabel("y")
plt.title(f"Order {order}")
plt.xlim((0, 1))
plt.ylim((-2, 2))
plt.tight_layout()
plt.show()
```

By switching the order we can see the new models generated. At first, using a larger order does better but after we reach a certain point, the model begins to overfit. Let's try to determine the best value for the hyperparemeter. In order to do that, we must select a metric to use to measure the effectiveness of a model. You can select any of the available [regression metrics from sklearn](http://scikit-learn.org/stable/modules/classes.html#regression-metrics) or develop your own.

In this example, we will start with the mean squared error. A better model will have a lower value for mean squared error.

```python
from sklearn.metrics import mean_squared_error
```

```python
orders = np.linspace(1,10,dtype=int)

scores = []
training_scores = []
for order in orders:
    polynomial_features = PolynomialFeatures(degree=order, include_bias=False)
    linear_regression = LinearRegression()
    pipeline = Pipeline([("polynomial_features", polynomial_features),
                             ("linear_regression", linear_regression)])
    pipeline.fit(X_model[:, np.newaxis], y_model)
    training_scores.append(mean_squared_error(y_model, pipeline.predict(X_model[:, np.newaxis])))
    scores.append(mean_squared_error(y_valid, pipeline.predict(X_valid[:, np.newaxis])))
```

```python
plt.plot(orders, training_scores, label="Training Score")
plt.plot(orders, scores, label="Validation Score")
plt.title("Model Comparison")
plt.xlabel("Order")
plt.ylabel("Mean Squared Error")
plt.legend(loc="best")
plt.show()
```

Now based on the above figure, select an order and let's see how it performs on the test set.

<!-- grade_id: selected_order, points: 0, grade: false, solution: true -->
```python
# Set the value of selectedOrder to what you think the best hyperparameter value should be
# YOUR CODE HERE
raise NotImplementedError()
```

```python
polynomial_features = PolynomialFeatures(degree=selectedOrder, include_bias=False)
linear_regression = LinearRegression()
pipeline = Pipeline([("polynomial_features", polynomial_features),
                         ("linear_regression", linear_regression)])

# Now we train on all the training data
pipeline.fit(X_train[:, np.newaxis], y_train)
# And test on the test data
test_score = mean_squared_error(y_test, pipeline.predict(X_test[:, np.newaxis]))
train_score = scores[selectedOrder-1]

print(f"The selected hyperparameter obtained a score of {train_score} in validation and a score of {test_score} in testing.")
if abs(test_score - train_score) < 0.01 or test_score < train_score:
    print("Good job!")
else:
    print("You might want to try again!")
```

<!-- grade_id: test_selected_order, points: 10, grade: true, solution: false -->
```python
assert abs(test_score - train_score) < 0.01 or test_score < train_score
```

## Feedback

<!-- grade_id: feedback, points: 0, grade: false, solution: true -->
```python
def feedback():
    """Provide feedback on the contents of this exercise
    
    Returns:
        string
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_feedback, points: 0, grade: true, solution: false -->
```python

```