## CSCI 470 Activities and Case Studies

1. For all activities, you are allowed to collaborate with a partner. 
1. For case studies, you should work individually and are **not** allowed to collaborate.

By filling out this notebook and submitting it, you acknowledge that you are aware of the above policies and are agreeing to comply with them.

Some considerations with regard to how these notebooks will be graded:

1. You can add more notebook cells or edit existing notebook cells other than "# YOUR CODE HERE" to test out or debug your code. We actually highly recommend you do so to gain a better understanding of what is happening. However, during grading, **these changes are ignored**. 
2. You must ensure that all your code for the particular task is available in the cells that say "# YOUR CODE HERE"
3. Every cell that says "# YOUR CODE HERE" is followed by a "raise NotImplementedError". You need to remove that line. During grading, if an error occurs then you will not receive points for your work in that section.
4. If your code passes the "assert" statements, then no output will result. If your code fails the "assert" statements, you will get an "AssertionError". Getting an assertion error means you will not receive points for that particular task.
5. If you edit the "assert" statements to make your code pass, they will still fail when they are graded since the "assert" statements will revert to the original. Make sure you don't edit the assert statements.
6. We may sometimes have "hidden" tests for grading. This means that passing the visible "assert" statements is not sufficient. The "assert" statements are there as a guide but you need to make sure you understand what you're required to do and ensure that you are doing it correctly. Passing the visible tests is necessary but not sufficient to get the grade for that cell.
7. When you are asked to define a function, make sure you **don't** use any variables outside of the parameters passed to the function. You can think of the parameters being passed to the function as a hint. Make sure you're using all of those variables.

# Unsupervised Learning - Clustering & Manifolds

In this exercise we will look at a variety of clustering methods including manifold-based methods.

First, we examine Agglomerative Clustering methods.

```python
import numpy as np
import scipy as sc
import pandas as pd
import matplotlib
import matplotlib.pyplot as plt
import sklearn as sk

%matplotlib inline
plt.style.use("ggplot")
```

```python
from sklearn.datasets import load_iris
from scipy.cluster.hierarchy import linkage, dendrogram
from sklearn.cluster import AgglomerativeClustering
from sklearn.metrics import adjusted_rand_score, adjusted_mutual_info_score, homogeneity_completeness_v_measure, silhouette_score
```

We will be using the [iris dataset](https://archive.ics.uci.edu/ml/datasets/iris) for this exercise.

```python
iris = load_iris()
iris_features = iris["data"]
iris_targets = iris["target"]
print(iris["DESCR"])
```

```python
iris_features.shape
```

```python
iris_targets.shape
```

## Agglomerative Clustering

In this portion of the exercise, you will plot the dendrograms when applied to the Iris dataset. To do this, you should use the [linkage](https://docs.scipy.org/doc/scipy/reference/generated/scipy.cluster.hierarchy.linkage.html#scipy.cluster.hierarchy.linkage) and [dendrogram](https://docs.scipy.org/doc/scipy-0.14.0/reference/generated/scipy.cluster.hierarchy.dendrogram.html#scipy.cluster.hierarchy.dendrogram) functions.

<!-- grade_id: cell-7ba61185fe878a6a, points: 0, grade: false, solution: true -->
```python
def dendrogram_plotter(features, methods, metric):
    """Plots a dendrogram for the provided features for every method in methods using the provided metric
    
    Args:
        features (iterable): The features to use in creating the dendrogram
        methods (iterable): A list of strings where each one is a valid method to the linkage function
        metric (str): A metric for calculating the linkage
    """
    for method in methods:
        plt.figure(figsize = (10,6)) # Change the figure size to your liking
        # YOUR CODE HERE
        raise NotImplementedError()
        plt.title(f"{method.title()} Linkage Iris Dataset Dendrogram")
        plt.show()
```

<!-- grade_id: cell-9a3884c6587cb41a, points: 5, grade: true, solution: false -->
```python
dendrogram_plotter(iris_features, ["average", "complete", "ward"], "euclidean")
```

<!-- grade_id: cell-1a45df2a94be8038, points: 0, grade: false, solution: true -->
```python
def clustering_scorer(features, targets, pred):
    """Calculates some important clustering scores given a set of features, targets, and predictions
    
    Args:
        features (iterable): The input features to the clustering problem
        targets (iterable): The targets if this was a classification problem
        pred (iterable): The cluster predictions for the data samples
    
    Returns:
        dict: A dictionary with the keys ['Adjusted Rand', 'Adjusted Mutual Info', 'Homogeneity', 'Completeness', 
            'V Score', 'Silhouette Score'] and values as the respective scores for the inputs.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

To create agglomerative clustering estimators, use [AgglomerativeClustering](http://scikit-learn.org/stable/modules/generated/sklearn.cluster.AgglomerativeClustering.html#sklearn.cluster.AgglomerativeClustering). It was already imported above. Consider using the function you just created above in developing the next function as well.

<!-- grade_id: cell-de19bff97996022b, points: 0, grade: false, solution: true -->
```python
def agg_clustering_scorer(features, targets, linkages, n_clusters=8):
    """Calculate the agglomerative clustering scores for a variety of linkage types
    
    Args:
        features (iterable): The input features of the data
        targets (iterable): The target classes if this was treated as a classification problem
        linkages (iterable): A list of linkage methods to calculate scores for
        n_clusters (int, optional): Defaults to 8. The number of clusters to use in the clustering algorithm
    
    Returns:
        iterable: Scores for each linkage method similar to the clustering_scorer method's output
    """
    scores = []
    for linkage in linkages:
        # YOUR CODE HERE
        raise NotImplementedError()
    return scores
```

```python
aggScores = agg_clustering_scorer(iris_features, iris_targets, ["average", "complete", "ward"], n_clusters=3)
```

<!-- grade_id: cell-8b0f9e6101ff9856, points: 5, grade: true, solution: false -->
```python
assert len(aggScores) == 3
for key in ['Adjusted Rand', 'Adjusted Mutual Info', 'Homogeneity', 'Completeness', 'V Score', 'Silhouette Score']:
    assert key in aggScores[0].keys()
```

```python
for linkage, score in zip(["average", "complete", "ward"],aggScores):
    print(f"With the {linkage} linkage,")
    print(f"Adjusted rand score is {score['Adjusted Rand']}")
    print(f"Adjusted mutual info score is {score['Adjusted Mutual Info']}")
    print(f"Homogeneity is {score['Homogeneity']}, Completeness is {score['Completeness']}, V score is {score['V Score']}")
    print(f"Silhouette score is {score['Silhouette Score']}\n")
```

```python
scoresdf = pd.DataFrame(aggScores, index=["average", "complete", "ward"]).T # Try removing the .T and see what happens
scoresdf.plot(kind="barh")
plt.title("Clustering scores with various linkage methods")
plt.show()
```

## Spectral Clustering

```python
from sklearn.cluster import SpectralClustering
```

<!-- grade_id: cell-fa37244cf09f0bae, points: 0, grade: false, solution: true -->
```python
# Create a spectral clustering classifier saved as `clf`, using n_clusters=3.
# Fit clf to the iris data
# Predict values for the iris data using clf and save them as spectral_pred
# YOUR CODE HERE
raise NotImplementedError()
```

```python
spectral_scores = clustering_scorer(iris_features, iris_targets, spectral_pred)
```

<!-- grade_id: cell-35596289ee109dc3, points: 5, grade: true, solution: false -->
```python
for key in ['Adjusted Rand', 'Adjusted Mutual Info', 'Homogeneity', 'Completeness', 'V Score', 'Silhouette Score']:
    assert key in spectral_scores.keys()
```

<!-- grade_id: cell-7144bc9fd9bb85a4, points: 5, grade: true, solution: false -->
```python
assert clf.n_clusters == 3
```

```python
if len(aggScores) == 3:
    aggScores.append(spectral_scores)
```

```python
scoresdf = pd.DataFrame(aggScores, index=["average", "complete", "ward", "spectral"]) # Try removing the .T and see what happens
scoresdf.plot(kind="barh")
plt.title("Clustering scores with various linkage methods and Spectral Clustering")
plt.show()
```

## Clustering Overview

The below code is copied from http://scikit-learn.org/stable/auto_examples/cluster/plot_cluster_comparison.html and is here for you to see how different methods work with varying input data. We haven't gone over all these methods and so you can focus only on what we've covered. We will go over some of the other methods in a later lecture.

```python
import time
import warnings

import numpy as np
import matplotlib.pyplot as plt

from sklearn import cluster, datasets, mixture
from sklearn.neighbors import kneighbors_graph
from sklearn.preprocessing import StandardScaler
from itertools import cycle, islice

np.random.seed(0)

# ============
# Generate datasets. We choose the size big enough to see the scalability
# of the algorithms, but not too big to avoid too long running times
# ============
n_samples = 1500
noisy_circles = datasets.make_circles(n_samples=n_samples, factor=.5,
                                      noise=.05)
noisy_moons = datasets.make_moons(n_samples=n_samples, noise=.05)
blobs = datasets.make_blobs(n_samples=n_samples, random_state=8)
no_structure = np.random.rand(n_samples, 2), None

# Anisotropicly distributed data
random_state = 170
X, y = datasets.make_blobs(n_samples=n_samples, random_state=random_state)
transformation = [[0.6, -0.6], [-0.4, 0.8]]
X_aniso = np.dot(X, transformation)
aniso = (X_aniso, y)

# blobs with varied variances
varied = datasets.make_blobs(n_samples=n_samples,
                             cluster_std=[1.0, 2.5, 0.5],
                             random_state=random_state)

# ============
# Set up cluster parameters
# ============
plt.figure(figsize=(9 * 2 + 3, 12.5))
plt.subplots_adjust(left=.02, right=.98, bottom=.001, top=.96, wspace=.05,
                    hspace=.01)

plot_num = 1

default_base = {'quantile': .3,
                'eps': .3,
                'damping': .9,
                'preference': -200,
                'n_neighbors': 10,
                'n_clusters': 3}

datasets = [
    (noisy_circles, {'damping': .77, 'preference': -240,
                     'quantile': .2, 'n_clusters': 2}),
    (noisy_moons, {'damping': .75, 'preference': -220, 'n_clusters': 2}),
    (varied, {'eps': .18, 'n_neighbors': 2}),
    (aniso, {'eps': .15, 'n_neighbors': 2}),
    (blobs, {}),
    (no_structure, {})]

for i_dataset, (dataset, algo_params) in enumerate(datasets):
    # update parameters with dataset-specific values
    params = default_base.copy()
    params.update(algo_params)

    X, y = dataset

    # normalize dataset for easier parameter selection
    X = StandardScaler().fit_transform(X)

    # estimate bandwidth for mean shift
    bandwidth = cluster.estimate_bandwidth(X, quantile=params['quantile'])

    # connectivity matrix for structured Ward
    connectivity = kneighbors_graph(
        X, n_neighbors=params['n_neighbors'], include_self=False)
    # make connectivity symmetric
    connectivity = 0.5 * (connectivity + connectivity.T)

    # ============
    # Create cluster objects
    # ============
    ms = cluster.MeanShift(bandwidth=bandwidth, bin_seeding=True)
    two_means = cluster.MiniBatchKMeans(n_clusters=params['n_clusters'])
    ward = cluster.AgglomerativeClustering(
        n_clusters=params['n_clusters'], linkage='ward',
        connectivity=connectivity)
    spectral = cluster.SpectralClustering(
        n_clusters=params['n_clusters'], eigen_solver='arpack',
        affinity="nearest_neighbors")
    dbscan = cluster.DBSCAN(eps=params['eps'])
    affinity_propagation = cluster.AffinityPropagation(
        damping=params['damping'], preference=params['preference'])
    average_linkage = cluster.AgglomerativeClustering(
        linkage="average", affinity="cityblock",
        n_clusters=params['n_clusters'], connectivity=connectivity)
    birch = cluster.Birch(n_clusters=params['n_clusters'])
    gmm = mixture.GaussianMixture(
        n_components=params['n_clusters'], covariance_type='full')

    clustering_algorithms = (
        ('MiniBatchKMeans', two_means),
        ('AffinityPropagation', affinity_propagation),
        ('MeanShift', ms),
        ('SpectralClustering', spectral),
        ('Ward', ward),
        ('AgglomerativeClustering', average_linkage),
        ('DBSCAN', dbscan),
        ('Birch', birch),
        ('GaussianMixture', gmm)
    )

    for name, algorithm in clustering_algorithms:
        t0 = time.time()

        # catch warnings related to kneighbors_graph
        with warnings.catch_warnings():
            warnings.filterwarnings(
                "ignore",
                message="the number of connected components of the " +
                "connectivity matrix is [0-9]{1,2}" +
                " > 1. Completing it to avoid stopping the tree early.",
                category=UserWarning)
            warnings.filterwarnings(
                "ignore",
                message="Graph is not fully connected, spectral embedding" +
                " may not work as expected.",
                category=UserWarning)
            algorithm.fit(X)

        t1 = time.time()
        if hasattr(algorithm, 'labels_'):
            y_pred = algorithm.labels_.astype(np.int)
        else:
            y_pred = algorithm.predict(X)

        plt.subplot(len(datasets), len(clustering_algorithms), plot_num)
        if i_dataset == 0:
            plt.title(name, size=18)

        colors = np.array(list(islice(cycle(['#377eb8', '#ff7f00', '#4daf4a',
                                             '#f781bf', '#a65628', '#984ea3',
                                             '#999999', '#e41a1c', '#dede00']),
                                      int(max(y_pred) + 1))))
        plt.scatter(X[:, 0], X[:, 1], s=10, color=colors[y_pred])

        plt.xlim(-2.5, 2.5)
        plt.ylim(-2.5, 2.5)
        plt.xticks(())
        plt.yticks(())
        plt.text(.99, .01, ('%.2fs' % (t1 - t0)).lstrip('0'),
                 transform=plt.gca().transAxes, size=15,
                 horizontalalignment='right')
        plot_num += 1

plt.show()
```

## Feedback

<!-- grade_id: feedback, points: 0, grade: false, solution: true -->
```python
def feedback():
    """Provide feedback on the contents of this exercise
    
    Returns:
        string
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: feedback-tests, points: 0, grade: true, solution: false -->
```python

```