## CSCI 470 Activities and Case Studies

1. For all activities, you are allowed to collaborate with a partner. 
1. For case studies, you should work individually and are **not** allowed to collaborate.

By filling out this notebook and submitting it, you acknowledge that you are aware of the above policies and are agreeing to comply with them.

# Optimization & Linear Algebra Review

In this exercise we will focus on the two main topics we covered in the lecture.

## Linear Algebra

In the lecture, with regards to Linear Algebra, we covered:

- some notation that we use for linear algebra
- special kinds of matrices
- matrix properties and operations
- norm definitions

Now let's practice some of these. 

```python
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
plt.style.use("ggplot")
```

In numpy we can create a vector or matrix or tensor using the [`np.array`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.array.html) object.

```python
vector = np.array([1,2,3])
print(vector)
print(f"The shape of our vector is {vector.shape}")
```

Output:
```
[1 2 3]
The shape of our vector is (3,)
```

```python
matrix = np.array([
    [1,2,3],
    [4,5,6]
])
print(matrix)
print(f"The shape of our matrix is {matrix.shape}")
```

Output:
```
[[1 2 3]
 [4 5 6]]
The shape of our matrix is (2, 3)
```

```python
tensor = np.array([
    [
        [1,2,3],
        [4,5,6]
    ],
    [
        [7,8,9],
        [10,11,12]
    ]
])
print(tensor)
print(f"The shape of our tensor is {tensor.shape}")
```

Output:
```
[[[ 1  2  3]
  [ 4  5  6]]

 [[ 7  8  9]
  [10 11 12]]]
The shape of our tensor is (2, 2, 3)
```

### Products

Let's try out some products now and check how they work. By default, numpy would do element wise multiplication if we use the `*` operator. Instead, we can use the [`@`](https://docs.python.org/3/library/operator.html#mapping-operators-to-functions) operator. (This is valid as of Python 3.5, otherwise you can use [`np.matmul`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.matmul.html)).

Let's start by making some vectors and matrices. 

```python
x = np.array([[1,2,3]]) # a row vector
y = np.array([[4],[5],[6]]) # a column vector
```

```python
print(x)
```

Output:
```
[[1 2 3]]
```

```python
print(y)
```

Output:
```
[[4]
 [5]
 [6]]
```

We can determine an array's dimensions by calling [`.shape`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.ndarray.shape.html).

```python
x.shape, y.shape
```

Output:
```
((1, 3), (3, 1))
```

Now let's try multiplying the two vectors to get their inner product. We'll also verify the result by manually calculating it.

```python
x @ y, 1 * 4 + 2 * 5 + 3 * 6
```

Output:
```
(array([[32]]), 32)
```

Now let's look at how transpose works. With numpy, you can access the transpose of a `np.array` by calling its [`.T`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.ndarray.T.html) method or using [`np.transpose()`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.transpose.html). For simplicity, we will use `.T`.

```python
print(x)
print(x.T)
```

Output:
```
[[1 2 3]]
[[1]
 [2]
 [3]]
```

In this next cell we can see that if the shapes of two different objects do not align, the product will fail.

```python
try:
    x.T @ y
except Exception as e:
    print(e)
    
try:
    x @ y.T
except Exception as e:
    print(e)
```

Output:
```
matmul: Input operand 1 has a mismatch in its core dimension 0, with gufunc signature (n?,k),(k,m?)->(n?,m?) (size 3 is different from 1)
matmul: Input operand 1 has a mismatch in its core dimension 0, with gufunc signature (n?,k),(k,m?)->(n?,m?) (size 1 is different from 3)
```

```python
x.shape, x.T.shape
```

Output:
```
((1, 3), (3, 1))
```

Squeezing gets rid of any shape dimensions with value = 1. This ensures a variable can be treated as a vector instead of a matrix with 1 row for example. Check out the documentation for [`.squeeze`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.squeeze.html).

```python
x.shape, x.squeeze().shape
```

Output:
```
((1, 3), (3,))
```

```python
x.T @ y.T
```

Output:
```
array([[ 4,  5,  6],
       [ 8, 10, 12],
       [12, 15, 18]])
```

Now let's move on to matrices.

```python
X = np.array([
    [1, 2],
    [3, 4]
])
Y = np.array([
    [1, 2, 3],
    [4, 5, 6]
])
X.shape, Y.shape
```

Output:
```
((2, 2), (2, 3))
```

<!-- grade_id: cell-613462977968b2ef, points: 0, grade: false, solution: true -->
```python
# Calculate the product of X and Y and save it to the variable `out`
# YOUR CODE HERE
out = X@Y
```

<!-- grade_id: cell-049e2d9010d235b1, points: 10, grade: true, solution: false -->
```python
assert out.shape == (2,3)
assert out[0,0] == 1 * 1 + 2 * 4
assert out[1,1] == 3 * 2 + 4 * 5
```

```python
print(f"The output shape of XY should be {X.shape[0]}, {Y.shape[1]}")
print(f"It is in fact, {out.shape} ")
```

Output:
```
The output shape of XY should be 2, 3
It is in fact, (2, 3)
```

### Norms

Now let's calculate the norms that we discussed in the lecture.

#### Vector Norms

First, let's calculate the $\ell_1$ norm. We can calculate norms manually or using numpy's [`linalg.norm`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.linalg.norm.html) method.

```python
ell_1_of_x = np.linalg.norm(x.squeeze(), 1)
sum_of_abs_of_x = np.abs(x).sum() 
print(ell_1_of_x, sum_of_abs_of_x)
```

Output:
```
6.0 6
```

Now calculate the $\ell_2$ norm of $\mathbf{x}$, once using the norm function and compare it to calculating it manually. Check the shape of $\mathbf{x}$ to make sure it is a vector. (You might need to use `.squeeze()`). 

<!-- grade_id: cell-699b991ec4a0d632, points: 0, grade: false, solution: true -->
```python
# Use the norm function to calculate ell_2_of_x 
# Manually calculate the norm and save to manual_ell_2_of_x

ell_2_of_x = np.linalg.norm(x)
manual_ell_2_of_x = x / x.sum(axis=1)[:, np.newaxis]

```

<!-- grade_id: cell-fff606eac0bb19f1, points: 5, grade: true, solution: false -->
```python
assert ell_2_of_x.all() == manual_ell_2_of_x.all()
```

#### Matrix Norms

Test the multiple definitions of the Frobenius norm.

```python
print(X)
result = []
for row in X:
    print(row)
    result.append(np.linalg.norm(row.flatten(), 2))
print(np.linalg.norm(result))
```

Output:
```
[[1 2]
 [3 4]]
[1 2]
[3 4]
5.477225575051661
```

```python
np.linalg.norm(X)
```

Output:
```
5.477225575051661
```

```python
print(np.sum(X**2)**0.5)
```

Output:
```
5.477225575051661
```

The trace of a matrix is the sum of its diagonal. With numpy we can calculate the trace of a matrix using [`np.trace`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.trace.html).

```python
print(X)
np.diag(X).sum(), np.trace(X)
```

Output:
```
[[1 2]
 [3 4]]
(5, 5)
```

Test the relationship between the Frobenius norm and trace.

```python
np.trace(X @ X.T)**0.5
```

Output:
```
5.477225575051661
```

### Special Matrices

Here we'll investigate the identity matrix, diagonal matrices, and symmetric matrices.

You can create an identity matrix using [`np.eye`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.eye.html)

```python
five_identity = np.eye(5)
five_identity
```

Output:
```
array([[1., 0., 0., 0., 0.],
       [0., 1., 0., 0., 0.],
       [0., 0., 1., 0., 0.],
       [0., 0., 0., 1., 0.],
       [0., 0., 0., 0., 1.]])
```

[`np.diag`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.diag.html) is used to either create a diagonal matrix or extract the diagonal of an existing matrix.

```python
D = np.diag([1,2,3,4,5])
D
```

Output:
```
array([[1, 0, 0, 0, 0],
       [0, 2, 0, 0, 0],
       [0, 0, 3, 0, 0],
       [0, 0, 0, 4, 0],
       [0, 0, 0, 0, 5]])
```

```python
np.diag(five_identity)
```

Output:
```
array([1., 1., 1., 1., 1.])
```

We can check if a matrix is Symmetric by checking if it is equivalent to its transpose. Is a diagonal matrix symmetric?


```python
print(f"It is {np.all(D == D.T)} that D is a symmetric matrix")
```

Output:
```
It is True that D is a symmetric matrix
```

```python
print(f"It is {np.all(X == X.T)} that X is a symmetric matrix")
```

Output:
```
It is False that X is a symmetric matrix
```

We can calculate the rank of a matrix using [`np.linalg.matrix_rank`](https://docs.scipy.org/doc/numpy-1.14.0/reference/generated/numpy.linalg.matrix_rank.html).

```python
np.linalg.matrix_rank(X)
```

Output:
```
2
```

```python
np.linalg.matrix_rank(five_identity)
```

Output:
```
5
```

```python
np.linalg.matrix_rank(np.array([
    [1, 2, 3],
    [2, 4, 6]
]))
```

Output:
```
1
```

```python
np.linalg.matrix_rank(np.array([
    [1, 2, 4],
    [2, 4, 6]
]))
```

Output:
```
2
```

Check to see if these numbers make sense to you based on the definition of rank.

## Decomposition

Now let's look at decomposing a matrix into multiple matrices. We have two main methods of doing so. Singular Value Decomposition works on all matrices and eigenvalue decomposition works on symmetric matrices.

We can use [`linalg.svd`](https://docs.scipy.org/doc/numpy-1.14.0/reference/generated/numpy.linalg.svd.html) to calculate the SVD of a matrix and [`linalg.eig`](https://docs.scipy.org/doc/numpy-1.14.0/reference/generated/numpy.linalg.eig.html) to form the eigenvalue decomposition of a symmetric matrix.

```python
U, Sigma, VT = np.linalg.svd(Y)
V = VT.T
print(U)
print(V)
print(Sigma)
U@U.T, U.T @ U, V@V.T, V.T @ V
```

Output:
```
[[-0.3863177  -0.92236578]
 [-0.92236578  0.3863177 ]]
[[-0.42866713  0.80596391  0.40824829]
 [-0.56630692  0.11238241 -0.81649658]
 [-0.7039467  -0.58119908  0.40824829]]
[9.508032   0.77286964]
(array([[ 1.00000000e+00, -8.49433239e-18],
        [-8.49433239e-18,  1.00000000e+00]]),
 array([[ 1.00000000e+00, -8.49433239e-18],
        [-8.49433239e-18,  1.00000000e+00]]),
 array([[1.00000000e+00, 1.16059657e-16, 9.83171390e-17],
        [1.16059657e-16, 1.00000000e+00, 4.58633852e-19],
        [9.83171390e-17, 4.58633852e-19, 1.00000000e+00]]),
 array([[ 1.00000000e+00,  4.13108370e-17,  2.21870637e-17],
        [ 4.13108370e-17,  1.00000000e+00, -1.77008767e-16],
        [ 2.21870637e-17, -1.77008767e-16,  1.00000000e+00]]))
```

Check the orthogonality of the resulting matrices.

```python
U@U.T, U.T @ U, V@V.T, V.T @ V
```

Output:
```
(array([[ 1.00000000e+00, -8.49433239e-18],
        [-8.49433239e-18,  1.00000000e+00]]),
 array([[ 1.00000000e+00, -8.49433239e-18],
        [-8.49433239e-18,  1.00000000e+00]]),
 array([[1.00000000e+00, 1.16059657e-16, 9.83171390e-17],
        [1.16059657e-16, 1.00000000e+00, 4.58633852e-19],
        [9.83171390e-17, 4.58633852e-19, 1.00000000e+00]]),
 array([[ 1.00000000e+00,  4.13108370e-17,  2.21870637e-17],
        [ 4.13108370e-17,  1.00000000e+00, -1.77008767e-16],
        [ 2.21870637e-17, -1.77008767e-16,  1.00000000e+00]]))
```

Let's create a symmetric matrix $\mathbf{S}$ to verify SVD's equivalence to eigenvalue decomposition when the matrix is symmetric.

```python
S = np.array([
    [1, 2, 3],
    [2, 3, 7],
    [3, 7, 12]
])
```

```python
U, Sigma, VT = np.linalg.svd(S)
V = VT.T
print(U)
print(V)
print(Sigma)
```

Output:
```
[[-0.22380127 -0.26826216  0.93698901]
 [-0.47060372  0.87162146  0.13714285]
 [-0.85348997 -0.41025777 -0.32131516]]
[[-0.22380127  0.26826216  0.93698901]
 [-0.47060372 -0.87162146  0.13714285]
 [-0.85348997  0.41025777 -0.32131516]]
[16.64636961  0.91033134  0.26396173]
```

<!-- grade_id: cell-da850a693bb7be9c, points: 0, grade: false, solution: true -->
```python
# Calculate the eigenvalues and eigenvectors of S and save them to eig_vals and eig_vecs respectively.
eig_vals, eig_vecs = np.linalg.eig(S)
```

<!-- grade_id: cell-0078b327366d23e9, points: 5, grade: true, solution: false -->
```python
assert eig_vals.shape == (3,)
assert eig_vecs.shape == (3,3)
```

Verify some properties of eigenvalues for symmetric matrices.

```python
np.linalg.eigvals(S.T @ S), np.linalg.eigvals(S @ S.T), np.linalg.eigvals(S) ** 2
```

Output:
```
(array([2.77101621e+02, 6.96757948e-02, 8.28703141e-01]),
 array([2.77101621e+02, 6.96757948e-02, 8.28703141e-01]),
 array([2.77101621e+02, 6.96757948e-02, 8.28703141e-01]))
```

## Feedback

<!-- grade_id: cell-9e98681da9a95117, points: 0, grade: false, solution: true -->
```python
def feedback():
    """Provide feedback on the contents of this exercise
    
    Returns:
        string
    """
    return "Good intro to Matrix and numpy cals. Would add something about why we use numpy and cases where we can optimize with numpy."
```

<!-- grade_id: cell-38e4606280755f60, points: 0, grade: true, solution: false -->
```python
feedback()
```

Output:
```
'Good intro to Matrix and numpy cals. Would add something about why we use numpy and cases where we can optimize with numpy.'
```

```python

```